Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "FANIS NIGAMADYANOV"
COLLABORATORS = ""
```

---

# Домашнее задание №4. Матричная факторизация

## Задачи - 25 баллов (+5 доп баллов)
1. Ноутбук `mf.ipynb` - 20 баллов
- ImplicitALS - 4 балла
- SVD - 4 балла
- Dataset with features - 2 балла
- ImplicitALS with features - 5 баллов
- LightFM with features - 5 баллов
2. Имплементация модели в сервис - 5 баллов
- Пробить на Leaderboard порог `map@10 = 0.075`
- Если при этом используете MF (Implicit или LightFM) + ANN (nmslib, faiss, annoy и тд) - дополнительно 5 баллов
  
## Как сдать ноутбук `mf.ipynb` на проверку

1. Прогоните весь код ноутбука - проверьте, что нет ошибок и тесты проходят
2. Выложите готовый ноутбук в ваш репозиторий с сервисом из домашнего задания №1 по пути `notebooks/hw_4/mf.ipynb` в ветке `hw_4`
3. Проверьте, что есть доступ к вашему репозиторию для аккаунтов `https://github.com/feldlime`
4. Откройте PR в main ветку и добавьте в ревьюеры **своего ментора**
5. Не проводите мердж в `main` ветку, пока не увидите оценку за это ДЗ в ведомости. Файл с ноутбуком должен находиться в ветке `hw_4`

Обратите внимание, что сборка ноутбуков на проверку автоматизирована. В случае неправильного пути, имени файла или ветки (а также при отсутствии доступа у `@feldlime`) ваша работа не попадёт на проверку и получит `0` баллов.

Используемые библиотеки в рамках ДЗ
```bash
pip install implicit==0.7.2 requests==2.32.3 rectools[lightfm]==0.12.0 pandas==2.2.3 numpy==1.26.4 scipy==1.12.0
```

## Импорты и данные

```python
import os
import threadpoolctl
import requests

import numpy as np
import pandas as pd
import zipfile as zf

import random

from tqdm.auto import tqdm
from implicit.als import AlternatingLeastSquares
from lightfm import LightFM

from rectools import Columns
from rectools.metrics import MAP, MeanInvUserFreq
from rectools.dataset import Dataset
from rectools.models import PureSVDModel, ImplicitALSWrapperModel, LightFMWrapperModel, model_from_config

import matplotlib.pyplot as plt
import seaborn as sns

# For implicit ALS
os.environ["OPENBLAS_NUM_THREADS"] = "1"
threadpoolctl.threadpool_limits(1, "blas")
```

Если вдруг у вас нет данных, то используйте закомментированный код

```python
# url = 'https://github.com/irsafilo/KION_DATASET/raw/f69775be31fa5779907cf0a92ddedb70037fb5ae/data_original.zip'

# req = requests.get(url, stream=True)

# with open('kion.zip', 'wb') as fd:
#     total_size_in_bytes = int(req.headers.get('Content-Length', 0))
#     progress_bar = tqdm(desc='kion dataset download', total=total_size_in_bytes, unit='iB', unit_scale=True)
#     for chunk in req.iter_content(chunk_size=2 ** 20):
#         progress_bar.update(len(chunk))
#         fd.write(chunk)

# files = zf.ZipFile('kion.zip', 'r')
# files.extractall()
# files.close()
```

```python
interactions = (
    pd.read_csv('data_original/interactions.csv', parse_dates=["last_watch_dt"])
    .rename(columns={'total_dur': Columns.Weight,
                     'last_watch_dt': Columns.Datetime})
)
users = pd.read_csv('data_original/users.csv')
items = pd.read_csv('data_original/items.csv')

user_features = users
user_features = user_features.rename(columns={'user_id': 'id'})
item_features = items
item_features = item_features.rename(columns={'item_id': 'id'})
```

```python
N_DAYS = 7

interactions['weight'] = np.log1p(interactions['weight'])

conditions = [
    interactions['watched_pct'].isna(),
    interactions['watched_pct'] <= 10,
    interactions['watched_pct'] > 10,
    interactions['watched_pct'] >= 50

]
choices = [1, 0.8, 1.2, 1.8]
interactions['weight'] = interactions['weight'] * np.select(conditions, choices)


max_date = interactions['datetime'].max()
train = interactions[(interactions['datetime'] <= max_date - pd.Timedelta(days=N_DAYS))]
test = interactions[(interactions['datetime'] > max_date - pd.Timedelta(days=N_DAYS))]

catalog = train[Columns.Item].unique()

test_users = test[Columns.User].unique()
users = train[Columns.User].unique()
cold_users = set(test_users) - set(train[Columns.User])
test.drop(test[test[Columns.User].isin(cold_users)].index, inplace=True)
hot_users = test[Columns.User].unique()

# dataset = Dataset.construct(train)
```

```python
K_RECOS = 10
map10 = MAP(k=K_RECOS)
```

## ImplicitALS

### Ситуация:

Коллега вернулся из отпуска и вы вместе сели за улучшение модели. Внимательно изучив репозиторий библиотеки implicit вы увидели модель iALS и решаете попробовать ее в деле.

Чтобы работа была интереснее, вы заключаете пари с вашим коллегой о том, кто выбьет больше MAP@K на горячих пользователях. 

Правила пари: 
- Валидируемся на последней неделе (переменная `test`) и на горячих пользователях `hot_users`
- Можно собрать свой `Dataset` на основе `train`, трансформированного, если нужно
- Параметры модели задаются конфигом, которые будут передаваться в `model_from_config`

У вашего коллеги получилось выбить на ImplicitALS `MAP@K = 0.052`. Ваша задача побить его рекорд.

<!-- grade_id: cell-9cc83595f9432cd2, points: 0, grade: false, solution: true -->
```python
# train preprocessing...
train2 = train.copy()
dataset = Dataset.construct(train2)
```

```python
config = {
    "cls": "ImplicitALSWrapperModel",
    "model": {
        "cls": "AlternatingLeastSquares",
        "factors": 2,
        "num_threads": 12,
        "iterations": 25,
        "random_state": 42,
        "regularization": 0.05,
        "alpha": 0.01,
    },
    "fit_features_together": True,
}
```

<!-- grade_id: cell-6c041ad93abfab9a, points: 4, grade: true, solution: false -->
```python
%%time
assert config['cls'] == 'ImplicitALSWrapperModel'

model = ImplicitALSWrapperModel.from_config(config)
model.fit(dataset)

recos = model.recommend(
    users=hot_users,
    dataset=dataset,
    k=K_RECOS,
    filter_viewed=True,
)

print(map10.calc(recos, test))

assert map10.calc(recos, test) >= 0.052
```

## SVD

На ваш громкий спор с коллегой о том, что все дело в вашем удачном random seed, к вам подошел ваш лид. 

Узнав детали вашего спора, он дает вам комментарий, что iALS хороша, но погружение в матричную факторизацию следует начинать с `SVD`.

Вы переглянулись с коллегой и решаете уладить спор о random seed во втором раунде, используя новую модель.

Ваш коллега смогу выбить на SVD `MAP@K = 0.066`. Вы знаете, что делать.

```python
dataset2 = Dataset.construct(train2)
print(dataset2)
```

<!-- grade_id: cell-f1b963c73db09e3a, points: 0, grade: false, solution: true -->
```python
config = {
    "cls": "PureSVDModel",
    "factors": 2,
    "maxiter": 25,
    "random_state": 42,
    "recommend_n_threads": 12,
}
```

<!-- grade_id: cell-366e2b318e3b502d, points: 4, grade: true, solution: false -->
```python
%%time
assert config['cls'] == 'PureSVDModel'

model = model_from_config(config)
model.fit(dataset2)

recos = model.recommend(
    users=hot_users,
    dataset=dataset2,
    k=K_RECOS,
    filter_viewed=True,
)
print(map10.calc(recos, test))

assert map10.calc(recos, test) >= 0.066
```

## Dataset with features

"Ну это ни в какие ворота!" - восклицает ваш коллега, увидев ваш победный конфиг. Из другого угла опенспейса доносится "А я говорил" от вашего лида.

В это время к вам сзади подходит продакт и интересуется предметом вашего спора.

Рассказав про особенности найденных вами моделей, он просит вас в них докинуть фичи, ведь на одних взаимодействиях далеко не уедешь.

Вы согласились, ведь это отличная возможность продолжить пари. Соберите `Dataset` с фичами по пользователям и итемам, который вы будете использовать дальше


```python
user_features
user_features_long = user_features.melt(
    id_vars=["id"],
    var_name="feature",
    value_name="value"
)
user_features_long
```

```python
def prepare_item_features(df: pd.DataFrame) -> pd.DataFrame:
    selected_cols = ["id", "content_type", "release_year", "genres", "countries", "age_rating"]
    
    df = df[selected_cols].copy()
    flat_features = []

    for _, row in df.iterrows():
        item_id = row["id"]
        for col in selected_cols[1:]:
            val = row[col]
            if pd.isna(val):
                continue

            if isinstance(val, str) and ',' in val:
                values = [v.strip() for v in val.split(',')]
            else:
                values = [val]
            for v in values:
                flat_features.append({
                    "id": item_id,
                    "feature": col,
                    "value": str(v)
                })

    return pd.DataFrame(flat_features)

item_features_long = prepare_item_features(item_features)
item_features_long
```

<!-- grade_id: cell-b0397a70b8235dd9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

dataset_with_features = Dataset.construct(
    interactions_df=train2,
    user_features_df=user_features_long,
    cat_user_features=["age", "income", "sex", "kids_flg"],
    item_features_df=item_features_long,
    cat_item_features = ["content_type", "release_year", "genres", "countries", "age_rating"]
)
```

<!-- grade_id: cell-7773b657b67af87b, points: 2, grade: true, solution: false -->
```python
assert (dataset_with_features.user_features is not None) and (dataset_with_features.item_features is not None)
```

## ImplicitALS with features

Собрав датасет с фичами вы готовы к третьему раунду пари. 

Вы решаете начать снова с `iALS`, до сих пор удивляясь результатам модели `SVD`.

Ваш коллега изучил вашу технику подбора random seed и хитро улыбается вам.

Он смог выбить `MAP@K = 0.073`, теперь ваш ход.

<!-- grade_id: cell-651def1b6c1a52b6, points: 0, grade: false, solution: true -->
```python
config = {
    "cls": "ImplicitALSWrapperModel",
    "model": {
        "cls": "AlternatingLeastSquares",
        "factors": 6,
        "num_threads": 12,
        "iterations": 25,
        "random_state": 42,
        "regularization": 0.05,
        "alpha": 0.01,
    },
    "fit_features_together": True,
}
```

<!-- grade_id: cell-f3ed67874a568322, points: 5, grade: true, solution: false -->
```python
%%time
assert config['cls'] == 'ImplicitALSWrapperModel'

model = model_from_config(config)
model.fit(dataset_with_features)

recos = model.recommend(
    users=hot_users,
    dataset=dataset_with_features,
    k=K_RECOS,
    filter_viewed=True,
)
print(map10.calc(recos, test))

assert map10.calc(recos, test) >= 0.073
```

## LightFM with features

И снова ор выше гор, ваш пайплайн подготовки датасета помог вам в очередной раз обойти вашего коллегу.

Не зная, к чему еще аппелировать, он зовет вашего старшего коллегу, чтобы тот внимательно изучил полученные результаты.

"iALS с фичами это хорошо, но тут стоит попробовать факторизационные машины, попробуйте `LightFM`" - заключает он. Вы переключаетесь на изучение новой библиотеки, предвкушая финальный раунд.

Ваш коллега смог выжать из своего обновленного `Dataset` и `LightFM` скор `MAP@10 = 0.08`. Последний рывок.

<!-- grade_id: cell-cec5f5df8bfd0aca, points: 0, grade: false, solution: true -->
```python
config = {
    "cls": "LightFMWrapperModel",
    "model": {
        "no_components": 16,
        "learning_rate": 0.05,
        "random_state": 42,
        "loss": "warp",
        "k": 3,
        "item_alpha": 0.05,
        "user_alpha": 0.05
    },
    "epochs": 4,
}

model = LightFMWrapperModel.from_config(config)
model.get_params(simple_types=True)
```

<!-- grade_id: cell-754f9152c2a343a5, points: 5, grade: true, solution: false -->
```python
%%time
assert config['cls'] == 'LightFMWrapperModel'

model = model_from_config(config)
model.fit(dataset_with_features)

recos = model.recommend(
    users=hot_users,
    dataset=dataset_with_features,
    k=K_RECOS,
    filter_viewed=True,
)
print(map10.calc(recos, test))

assert map10.calc(recos, test) >= 0.08
```

## Сервис

Эта битва была легендарной, но вот за спиной опять появляется ваш продакт. 

"Время катить АБ" - говорит он в пятницу вечером. Делать нечего, вы собираете ваши наработки и определяете совместный фронт работ.

В прод должна заехать лучшая модель, которая побьет текущей модели в проде в `MAP@10 = 0.075`.

Также есть бонус от вашего старшего коллеги, который вам советует присмотреться к `Approximate nearest neighbours`, например к `nmslib`.

"Если сможешь обернуть в ANN, то на следующем годовом ревью получишь от меня оценку отлично" - сказал он. Изучить новые технологии и получить повышение - идеально, заключаете вы и бросаетесь в бой.

# Вместо заключения

## Задачи - 25 баллов (30 баллов с доп задачей по ANN)
1. Ноутбук `mf.ipynb` - 20 баллов
- SVD - 5 баллов
- ImplicitALS - 5 баллов
- ImplicitALS with features - 5 баллов
- LightFM with features - 5 баллов
2. Имплементация модели в сервис - 5 баллов
- Пробить на Leaderboard порог `map@10 = 0.075`
- Если при этом используете MF (Implicit или LightFM) + ANN (nmslib, faiss, annoy и тд) - дополнительно 5 баллов
  
## Как сдать ноутбук `mf.ipynb` на проверку

1. Прогоните весь код ноутбука - проверьте, что нет ошибок и тесты проходят
2. Выложите готовый ноутбук в ваш репозиторий с сервисом из домашнего задания №1 по пути `notebooks/hw_4/mf.ipynb` в ветке `hw_4`
3. Проверьте, что есть доступ к вашему репозиторию для аккаунтов `https://github.com/feldlime`
4. Откройте PR в main ветку и добавьте в ревьюеры **своего ментора**
5. Не проводите мердж в `main` ветку, пока не увидите оценку за это ДЗ в ведомости. Файл с ноутбуком должен находиться в ветке `hw_4`

Обратите внимание, что сборка ноутбуков на проверку автоматизирована. В случае неправильного пути, имени файла или ветки (а также при отсутствии доступа у `@feldlime`) ваша работа не попадёт на проверку и получит `0` баллов.